<a href="https://colab.research.google.com/github/mhpolas/-Dashboard-using-Dash-and-Plotly-/blob/main/4_dicts_classes.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Fill in your github username and idm credentials below:

```python
github_username = "od41atam"
idm_credentials = "mhpolas"
```

---

# **Applied Data Science in Medicine & Psychology - Assignment 4**
## Dictionaries and Classes
---
### **General Setup and Submission**
1. Save an editable copy of this notebook to your Google Drive ``File`` $\rightarrow$ ``Save a copy in Drive`` and remove the automatically appended "Copy of" text. **Do not change the original file name**.
2. Work on the assignment in the *copied* version of the notebook. Make sure that you **replace** all parts of the tasks indicated with *YOUR CODE HERE* and **raise NotImplementedError()** (otherwise this causes an error in your implementation) or *YOUR ANSWER HERE* with your solution.
3. If you take a break during the implementation, you can save the current status of your work in your repository following the submission instructions below. When you continue working on the assignment, do this in the version of the notebook that you saved in your Google Drive (access from your Google Colab account) or your repository.
4. After finishing the assignment submit your assignment as follows:
    * Check that everything runs as expected. To do that select in the menubar ``Runtime`` $\rightarrow$ ``Restart and run all``. This will clear all your local variables and runs your notebook from the beginning.
    * Save your notebook, click ``File`` $\rightarrow$ ``Save``
    * Click ``File`` $\rightarrow$ ``Save a copy in GitHub``
    * Select ``digikolleg-data-science-psychology/4-dicts-classes-<username>`` from the repository drop-down menu
    * Type ``4_dicts_classes.ipynb`` into the file path text box
    * Optionally, you may enter a description into the "Commit message" text box (like: "solved task 1" or "ready for grading")
    * Make sure that the "Include a link to Colaboratory" box is checked
    * Press "Ok"
    * A new tab or window should open up, displaying your submitted notebook within your GitHub repository

Before the assignment deadline, you can re-submit your assignment as often as you would like. This will update your assignment file in your assignment's repository.

### **Assignment Description**
In this assignment, you will explore two essential concepts in Python. *Dictionaries* are data structures that allow you to store, retrieve, and group data. *Classes* are the building blocks of object-oriented programming (OOP) and provide a way to create reusable, modular code. Learning *dictionaries* and *classes* fosters an object-oriented mindset, which is a powerful and widely-used programming paradigm.

#### **Assignment Goals**
You are going to learn the following things in this assignment:
1. how to create a dictionary and add, remove, and change its elements
2. how to create classes and instances of classes
3. how to write class-specific methods
4. how to use class inheritance
5. how to define relationships between objects

#### **Assignment Point Value**
This assignment is worth a total (up to) **13** points. If your code passes all of the manual and automatic tests you will receive full credit.

### **Tasks**
Make sure that you now work on your *copied* version of this assignment's notebook.

#### 0. Imports and Setup
Please **run** the cell below to import the libraries needed for this assignment.

```python
from datetime import date, datetime
from contextlib import suppress
```

#### 1. Dictionary
A dictionary is a type of data structure that stores key-value pairs. Each key in the dictionary is associated with a value:
```
dict_1 = {
    key_1: value_1,
    key_2: value_2
}
```

##### 1.1 Create a dictionary called ``patient`` with the following key-value pairs:

| key | value |
| :- | :- |
|first_name|Jane|
|last_name|Doe|
|height|167.3|
|vegetarian|True|
|allergies|[pollen, dust, peanuts]|
|blood_test_results|None|

<!-- grade_id: create_dict, points: 0, grade: false, solution: true -->
```python
patient ={"first_name": "Jane",
          "last_name": "Doe",
          "height": 167.3,
          "vegetarian": True,
          "allergies": ["pollen", "dust", "peanuts"],
          "blood_test_results": None,}

```

<!-- grade_id: create_dict_test, points: 0.5, grade: true, solution: false -->
```python
# run tests to check your implementation
assert len(patient) == 6, "Check if you added all key-value pairs"
assert list(patient.keys()) == ['first_name', 'last_name', 'height', 'vegetarian', 'allergies', 'blood_test_results'], "Check if you added all key-value pairs correctly"
assert patient["first_name"] == "Jane", "Check if you assigned the correct value to the keys"
```

##### 1.2 Add the key ``age`` with value 24 to the ``patient`` dictionary:

<!-- grade_id: add_key, points: 0, grade: false, solution: true -->
```python
patient["age"] = 24
```

<!-- grade_id: add_key_test, points: 0.5, grade: true, solution: false -->
```python
# run test to check your implementation
assert patient["age"], "Check if you added the key-value pair correctly"
```

##### 1.3 Create another dictionary called ``blood_test_results`` and change the ``blood_test_results`` key of ``patient``:
| key | value |
| :- | :- |
|hemoglobin|13.6|
|potassium|4.5|

<!-- grade_id: change_dict, points: 0, grade: false, solution: true -->
```python
blood_test_results = {"hemoglobin": 13.6,
                       "potassium": 4.5}
patient["blood_test_results"] = blood_test_results
```

<!-- grade_id: change_dict_test, points: 0.5, grade: true, solution: false -->
```python
# run test to check your implementation
assert len(patient["blood_test_results"]) == 2, "Check if you added both elements to the dictionary"
```

##### 1.4 Remove the key-value pair associated with ``allergies`` from ``patient``:

<!-- grade_id: remove_key, points: 0, grade: false, solution: true -->
```python
patient.pop("allergies")
print(patient)
```

Output:
```
{'first_name': 'Jane', 'last_name': 'Doe', 'height': 167.3, 'vegetarian': True, 'blood_test_results': {'hemoglobin': 13.6, 'potassium': 4.5}, 'age': 24}
```

<!-- grade_id: remove_key_test, points: 0.5, grade: true, solution: false -->
```python
# run test to check your implementation
with suppress(KeyError): patient["allergies"]; assert False, "Ensure that you removed the correct key-value pair"
```

#### 2. Classes - Medical Records System
In this task, you will create a simulation of a patient's medical records system. Therefore, you will create 5 different classes:
* ``Medication``: Stores information like name and dosage for one type of medication. It is possible to change the prescribed dosage.
* ``Session``: Contains the information about any appointment a patient could have. Also, notes can be stored after the session.
* ``Physician``: Is the base class for every physician. A physician can schedule sessions for their patients and add notes to the sessions.
* ``PrimaryCarePhysician``: Inherits all functionalities from ``Physician`` and is able to refer a patient to a specialist.
* ``Patient``: Is the base of the medical records system. All medications, sessions, and referrals are stored there. Also, any modifications are initiated using this class. Furthermore, it can visualize all stored entries for the user.

##### 2.1 Create the class ``Medication``, to do so:
* Create a class ``Medication``
* Create the constructor (``__init__()``) with ``name`` and ``dosage`` as parameters. Initialize the attributes of the object ``self.name`` and ``self.dosage`` with the respective parameter.
* Create the method ``update_dosage`` with ``new_dosage`` as parameter which updates the object attribute ``self.dosage``
* Provide a meaningful string representation of your object using the ``__str__(self)`` method (one of the [special methods](https://rszalski.github.io/magicmethods/#representations) in Python). For example, if the object's ``name`` is "Ibuprofen" and ``dosage`` is "200 mg" you could return: "Medication: Ibuprofen, Dosage: 200 mg". You can use [Formatted String Literals](https://docs.python.org/3/tutorial/inputoutput.html#tut-f-strings) for that.
* Create an instance of ``Medication`` to test your implementation and try to access its object attributes

**Hint**: If the instantiation doesn't work properly, have a look at the ``self`` keyword and its meaning and usage.

<!-- grade_id: medication, points: 0, grade: false, solution: true -->
```python
class Medication:
    def __init__(self, name, dosage):
        self.name = name
        self.dosage = dosage

    def update_dosage(self, new_dosage):
        self.dosage = new_dosage

    def __str__(self):
        return f"Medication: {self.name}, Dosage: {self.dosage}"
```

<!-- grade_id: medication_test, points: 1, grade: true, solution: false -->
```python
# run tests to check your implementation
medication = Medication("Ibuprofen", "200 mg")
assert medication.name == "Ibuprofen" and  medication.dosage == "200 mg", "Ensure that you save the constructor parameters as object attributes"
medication.update_dosage("400 mg"),
assert medication.dosage == "400 mg", "Ensure that you save new_dosage in the respective object attribute"
assert medication.name in str(medication) and medication.dosage in str(medication), "Ensure that you return a string with all the relevant information"
```

##### 2.2 Create the class ``Session``, to do so:
* Create a class ``Session``
* Create a constructor with ``session_date``, ``physician_name``, and ``purpose`` as parameters
    * Save ``session_date`` only as object attribute if it is of type ``datetime.date``. Otherwise, raise a TypeError with a meaningful message for the user. The class ``date`` is already imported from the module ``datetime`` at the beginning of the assignment.
    * Save ``physician_name`` only as object attribute if it is of type ``str``. Otherwise raise a TypeError with a meaningful message.
    * Save ``purpose`` as an object attribute with the name accordingly
    * Furthermore, create an object attribute ``notes`` and initialize it with an empty dictionary.
* Create a method ``add_session_notes`` with the parameter ``notes``. Save ``notes`` in the respective object attribute using the current local date and time as key. Check out the [datetime](https://docs.python.org/3/library/datetime.html#module-datetime) module for a suitable function.
* Provide a string representation of ``Session`` like for ``Medication``. Include all object attributes (also the notes). **Hint**: ``date`` also has a built-in string representation.
* Create an instance of ``Session`` and test your methods

<!-- grade_id: session, points: 0, grade: false, solution: true -->
```python
class Session:
    def __init__(self, session_date, physician_name, purpose):
        if isinstance(session_date, date):
            self.session_date = session_date
        else:
            raise TypeError("session_date must be a datetime.date object")


        if isinstance(physician_name, str):
            self.physician_name = physician_name
        else:
            raise TypeError("physician_name must be a string")


        self.purpose = purpose
        self.notes = {}

    def add_session_notes(self, notes):

        self.notes[datetime.now()] = notes


    def __str__(self):
        notes_str = ", ".join([f"{k.strftime('%Y-%m-%d %H:%M:%S')}: {v}" for k, v in self.notes.items()])
        return f"Session Date: {self.session_date}, Physician: {self.physician_name}, Purpose: {self.purpose}, Notes: {{{notes_str}}}"
```

<!-- grade_id: session_test, points: 2, grade: true, solution: false -->
```python
# run tests to check your implementation
session = Session(date(year=2000, month=1, day=1), "Dr. X", "blood_test")
assert session.session_date == date(year=2000, month=1, day=1) and session.physician_name == "Dr. X" and session.purpose == "blood_test" and len(session.notes) == 0, "Ensure that you set all object attributes correctly"
with suppress(TypeError): Session("2000, 1, 1", None, None); assert False, "Ensure that you raise the correct type of Exception"
with suppress(TypeError): Session(date(year=2000, month=1, day=1), None, None); assert False, "Ensure that you raise the correct type of Exception"
session.add_session_notes("anxiety")
assert set(map(type, session.notes)) == {datetime} and set(session.notes.values()) == {"anxiety"}, "Ensure that you used datetime as key and the given notes as values"
assert str(session.session_date) in str(session) and session.physician_name in str(session) and session.purpose in str(session) and "anxiety" in str(session), "Ensure that you use all relevant information in your string representation"
```

##### 2.3 Create the class ``Patient`` with all its three subparts **medication handling**, **session handling**, and **referral handling**, to do so:
* Create a class ``Patient``
* Create the constructor which takes ``name`` as parameter and saves it as an object attribute
    * Initialize object attributes ``medications``, ``sessions``, and ``referrals`` as empty dictionaries
    * Initialize an object attribute ``session_id`` with 0

**Medication Handling:**
* Create a method ``add_medication`` which takes an object of ``Medication`` as parameter ``medication``
    * Add ``medication`` to the respective object attribute using its name as key
* Create a method ``remove_medication`` with ``medication_name`` as parameter
    * Remove the medication from the respective object attribute
* Create a method ``update_medication_dosage`` with ``medication_name`` and ``new_dosage`` as parameters
    * Update the dosage of the given medication
* Create a method ``print_medications`` which prints all the prescribed medications for this patient using the defined string representation of ``Medication``

**Session Handling:**
* Create a methode ``increment_session_id`` which increments ``session_id`` by 1
* Create a method ``add_therapy_session`` with ``session`` as parameter
    * Perform a type check of ``session``. If it is of type ``Session`` add ``session`` to the respective object attribute using ``session_id``. Otherwise, raise a TypeError with a meaningful message.
    * Note: ``session_id`` must be unique, first session gets an id equal to 0
* Create a method ``get_session`` with ``session_id`` as parameter which returns the corresponding ``Session`` object
* Create a method ``print_sessions`` which prints all sessions for this patient using the defined string representation of ``Session``

**Referral Handling:**
* Create a method ``refer_to_specialist`` with ``specialist`` (of type str) as parameter
    * Add ``specialist`` to the respective object attribute using the current local date and time as key
* Create a method ``print_referrals`` which prints all referrals for this patient

**Testing:**
* Create an instance of patient
* Use the different methods to test your implementation
* Try your ``print_...`` methods. Do they show all relevant information? If not check your string representations of ``Medication`` and ``Session``.

<!-- grade_id: patient, points: 0, grade: false, solution: true -->
```python
class Patient:
    def __init__(self, name):
        self.name = name
        self.medications = {}
        self.sessions = {}
        self.referrals = {}
        self.session_id = 0

    def add_medication(self, medication):
        """Adds a Medication object to the patient's medications."""
        self.medications[medication.name] = medication

    def remove_medication(self, medication_name):
        """Removes a medication by its name."""
        self.medications.pop(medication_name, None)

    def update_medication_dosage(self, medication_name, new_dosage):
        """Updates the dosage of a specific medication."""
        if medication_name in self.medications:
            self.medications[medication_name].update_dosage(new_dosage)
        else:
            print(f"Medication '{medication_name}' not found for {self.name}.")

    def print_medications(self):
        """Prints all prescribed medications for the patient."""
        print(f"Medications for {self.name}:")
        if not self.medications:
            print("No medications prescribed.")
        else:
            for medication in self.medications.values():
                print(medication)


    def increment_session_id(self):
        """Increments the session ID."""
        self.session_id += 1

    def add_therapy_session(self, session):
        """Adds a Session object to the patient's sessions."""
        if isinstance(session, Session):
            self.increment_session_id()
            self.sessions[self.session_id] = session
        else:
            raise TypeError("session must be a Session object")

    def get_session(self, session_id):
        """Returns the Session object for a given session ID."""
        return self.sessions.get(session_id)

    def print_sessions(self):
        """Prints all sessions for the patient."""
        print(f"Sessions for {self.name}:")
        if not self.sessions:
            print("No sessions recorded.")
        else:

            for session_id in sorted(self.sessions.keys()):
                print(f"Session ID: {session_id}, {self.sessions[session_id]}")



    def refer_to_specialist(self, specialist):
        """Adds a referral to a specialist for the patient."""
        self.referrals[datetime.now()] = specialist

    def print_referrals(self):
        """Prints all referrals for the patient."""
        print(f"Referrals for {self.name}:")
        if not self.referrals:
            print("No referrals recorded.")
        else:

            for referral_time in sorted(self.referrals.keys()):
                print(f"{referral_time.strftime('%Y-%m-%d %H:%M:%S')}: {self.referrals[referral_time]}")
```

Run tests to check your implementation for **medication handling**:

<!-- grade_id: patient_medication_test, points: 2, grade: true, solution: false -->
```python
patient = Patient("Jane Doe")
assert patient.name == "Jane Doe" and len(patient.medications) == 0 and len(patient.sessions) == 0 and len(patient.referrals) == 0 and patient.session_id == 0, "Ensure that you set all object attributes correctly"
med = Medication("Ibuprofen", "200 mg")
patient.add_medication(med)
assert len(patient.medications) == 1 and set(patient.medications.keys()) == {med.name} and patient.medications[med.name] == med, "Ensure that you used the medication name as key and the Medication object as value"
patient.update_medication_dosage(med.name, "400 mg")
assert len(patient.medications) == 1 and patient.medications[med.name].dosage == "400 mg", "Ensure that you used the Medication method update_dosage"
patient.remove_medication(med.name)
assert len(patient.medications) == 0, "Check if you delete the Medication object in remove_medication"
```

Run tests to check your implementation for **session handling**:

<!-- grade_id: patient_session_test, points: 2, grade: true, solution: false -->
```python
patient = Patient("Jane Doe")
patient.increment_session_id()
assert patient.session_id == 1, "Check if your stored the incremented session_id as object attribute"
session = Session(date(year=2000, month=1, day=1), "Dr. X", "blood_test")
patient.add_therapy_session(session)
assert len(patient.sessions) == 1 and set(patient.sessions.keys()) == {1} and patient.sessions[1] == session, "Ensure that you added the session correctly in add_therapy_session"
assert patient.session_id == 2, "Check if you incremented session_id after session adding"
assert patient.get_session(1) == session, "Ensure that you return the correct Session object"
```

Run tests to check your implementation for **referral handling**:

<!-- grade_id: referral_test, points: 1, grade: true, solution: false -->
```python
patient = Patient("Jane Doe")
patient.refer_to_specialist("psychologist")
assert len(patient.referrals) == 1 and set(map(type, patient.referrals.keys())) == {datetime} and set(patient.referrals.values()) == {"psychologist"}, "Ensure that you added the referral correctly"
```

##### 2.4 Create the class ``Physician``, to do so:
* Create a class ``Physician``
* Create the constructor which takes ``name`` as parameter and stores it as an object attribute
* Create the method ``schedule_session`` with ``session_date``, ``patient``, and ``purpose`` as parameters
    * Create an instance of ``Session`` with the given information
    * Add the session to the therapy sessions of ``patient``
* Create the method ``add_session_notes`` with ``patient``, ``session_id``, and ``notes`` as parameters
    * Use the ``get_session`` method of ``patient`` to get the respective ``Session`` object of ``session_id``
    * add ``notes`` to the session

<!-- grade_id: physician, points: 0, grade: false, solution: true -->
```python
class Patient:
    def __init__(self, name):
        self.name = name
        self.medications = {}
        self.sessions = {}
        self.referrals = {}
        self.session_id = 0


    def add_medication(self, medication):
        """Adds a Medication object to the patient's medications."""
        self.medications[medication.name] = medication

    def remove_medication(self, medication_name):
        """Removes a medication by its name."""
        self.medications.pop(medication_name, None)

    def update_medication_dosage(self, medication_name, new_dosage):
        """Updates the dosage of a specific medication."""
        if medication_name in self.medications:
            self.medications[medication_name].update_dosage(new_dosage)
        else:
            print(f"Medication '{medication_name}' not found for {self.name}.")

    def print_medications(self):
        """Prints all prescribed medications for the patient."""
        print(f"Medications for {self.name}:")
        if not self.medications:
            print("No medications prescribed.")
        else:
            for medication in self.medications.values():
                print(medication)

    def increment_session_id(self):
        """Increments the session ID."""
        self.session_id += 1

    def add_therapy_session(self, session):
        """Adds a Session object to the patient's sessions."""
        if isinstance(session, Session):

            self.sessions[self.session_id] = session

            self.increment_session_id()
        else:
            raise TypeError("session must be a Session object")

    def get_session(self, session_id):
        """Returns the Session object for a given session ID."""
        return self.sessions.get(session_id)

    def print_sessions(self):
        """Prints all sessions for the patient."""
        print(f"Sessions for {self.name}:")
        if not self.sessions:
            print("No sessions recorded.")
        else:

            for session_id in sorted(self.sessions.keys()):
                print(f"Session ID: {session_id}, {self.sessions[session_id]}")



    def refer_to_specialist(self, specialist):
        """Adds a referral to a specialist for the patient."""
        self.referrals[datetime.now()] = specialist

    def print_referrals(self):
        """Prints all referrals for the patient."""
        print(f"Referrals for {self.name}:")
        if not self.referrals:
            print("No referrals recorded.")
        else:

            for referral_time in sorted(self.referrals.keys()):
                print(f"{referral_time.strftime('%Y-%m-%d %H:%M:%S')}: {self.referrals[referral_time]}")
```

<!-- grade_id: physician_test, points: 2, grade: true, solution: false -->
```python
# run tests to check your implementation
physician = Physician("Dr. X")
assert physician.name == "Dr. X"
patient = Patient("Jane Doe")
physician.schedule_session(date(2000, 1, 1), patient, "blood_test")
assert len(patient.sessions) == 1, "Ensure that you use the patient's add_therapy_session method"
physician.add_session_notes(patient, 0, "patient reported anxiety")
session = patient.get_session(0)
assert set(session.notes.values()) == {"patient reported anxiety"}, "Ensure that you used the add_session_notes of session"
```

##### 2.5 Create the class ``PrimaryCarePhysician``, to do so:
* Create a class ``PrimaryCarePhysician`` that inherits from ``Physician``
* Create a method ``refer_to_specialist`` with ``patient`` and ``specialist`` (type str) as parameters
    * Use the ``refer_to_specialist`` method from ``Patient`` to add the referral to ``patient``

<!-- grade_id: prim_care_physician, points: 0, grade: false, solution: true -->
```python
class PrimaryCarePhysician(Physician):
    def refer_to_specialist(self, patient, specialist):
        patient.refer_to_specialist(specialist)
```

<!-- grade_id: prim_care_physician_test, points: 1, grade: true, solution: false -->
```python
# run tests to check your implementation
prim_care_phys = PrimaryCarePhysician("Dr. Z")
patient = Patient("Jane Doe")
prim_care_phys.schedule_session(date(2000, 1, 1), patient, "blood_test")
assert len(patient.sessions) == 1, "Ensure that PrimaryCarePhysician inherits all properties from Physician"
prim_care_phys.refer_to_specialist(patient, "psychologist")
assert len(patient.referrals) == 1, "Ensure that you used the patient's refer_to_specialist method"
```

##### 2.6 Play around with the patient's medical records system that you created:
* Create one or more patients
* Create different physicians and schedule sessions for your patients
* Add the session notes
* Add, remove, and change your patients' medications
* refer your patients to different specialists
* Use your patients' visualization system (``print_medications()``... methods) to keep track of all the changes

<!-- grade_id: playground, points: 0, grade: false, solution: true -->
```python
patient1 = Patient("Alice")
patient2 = Patient("Bob")
```

```python

```